# Problem 1: Converting math to Python (*3 points*)

One of the goals of this part of the course is to develop your quantitative geoscience skills, including learning how to convert mathematical equations to Python code. Doing this allows you to explore how various equations work and produce useful data plots or predictions, something increasingly done by geoscience professionals.

For this problem you are asked to create 3 Python functions to calculate common statistics on a list of values: (1) the *mean*, (2) the *standard deviation*, and (3) the *standard deviation of the mean* (or *standard error*). In Problem 2 we'll use these functions to calculate some statistics related to Holocene volcanoes from the [Smithsonian Institution's Global Volcanism Program](https://volcano.si.edu/).

**Notice**: Closely follow the instructions! For example, you should be sure to use **exactly** the same variable names mentioned in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions.

**Your score on this problem will be based on following criteria**:

- Properly converting the given formulas to Python functions
    - **Note**: You should not use existing Python, Pandas or NumPy functions in the functions you create, other than perhaps a function for calculating the square root
- Including comments that explain what most lines in the code do
- Uploading your notebook to your GitHub repository for this week's exercise

## Part 1: Creating your `mean()` function (*1 point*)

The *mean* or *average* $\bar{x}$ should be calculated using a function you should call `mean()`.
  
\begin{equation}
  \Large
  \bar{x} = \frac{\Sigma x_{i}}{N}
\end{equation}

*Equation 1. The mean value, where $x_{i}$ is a value to be included in the mean calculation and $N$ is the total number of values to average*.

- Create a function called `mean()` in the cell below.

<!-- grade_id: mean, points: 0, grade: false, solution: true -->
```python
def mean(numbers):
    """Returns the average value of a collection of numbers."""
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_1_mean_tests, points: 1, grade: true, solution: false -->
```python
# These tests with small lists should work

from nose.tools import assert_equal, assert_true

# Test mean function with some lists
list1 = [1, 2, 3, 4, 5]
list2 = [-1, 0, -5, 2]

# Calculate list means
list1_mean = mean(list1)
list2_mean = mean(list2)

# Print mean values
print("Mean for list1:", list1_mean)
print("Mean for list2:", list2_mean)

# Check that the mean values are correct
assert_equal(round(list1_mean, 3), 3.000)
assert_equal(round(list2_mean, 3), -1.000)

```

## Part 2: Creating your `stddev()` function (*1 point*)

The *standard deviation* $\sigma_{x}$, calculated using a function you should call `stddev()`.

\begin{equation}
  \Large
  \sigma_{x} = \sqrt{\frac{1}{N} \Sigma \left( x_{i} - \bar{x} \right)^{2}}
\end{equation}

*Equation 2. The standard deviation*.

- Create a function called `stddev()` in the cell below.

<!-- grade_id: stddev, points: 0, grade: false, solution: true -->
```python
# Import NumPy to use the square root function math.sqrt()
import math

def stddev(numbers):
    """Returns the standard deviation of a collection of numbers."""
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_1_stddev_tests, points: 1, grade: true, solution: false -->
```python
# These tests with small lists should work

from nose.tools import assert_equal, assert_true

# Test stddev function with some lists
list1 = [1, 2, 3, 4, 5]
list2 = [-1, 0, -5, 2]

# Calculate list stddev
list1_stddev = stddev(list1)
list2_stddev = stddev(list2)

# Print stddev values
print("Standard deviation for list1:", list1_stddev)
print("Standard deviation for list2:", list2_stddev)

# Check that the stddev values are correct
assert_equal(round(list1_stddev, 3), 1.414)
assert_equal(round(list2_stddev, 3), 2.550)

```

## Part 3: Creating your `stderr()` function (*1 point*)

The *standard deviation of the mean* or *standard error* $\sigma_{\bar{x}}$, calculated using a function you should call `stderr()`.

\begin{equation}
  \Large
  \sigma_{\bar{x}} = \frac{\sigma_{x}}{\sqrt{N}}
\end{equation}

*Equation 3. The standard error*.

- Create a function called `stderr()` in the cell below.

<!-- grade_id: stderr, points: 0, grade: false, solution: true -->
```python
# Import the math module to use the square root function math.sqrt()
import math

def stderr(numbers):
    """Returns the standard error of a collection of numbers."""
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: stderr_teproblem_1_stderr_testssts, points: 1, grade: true, solution: false -->
```python
# These tests with small lists should work

from nose.tools import assert_equal, assert_true

# Test stderr function with some lists
list1 = [1, 2, 3, 4, 5]
list2 = [-1, 0, -5, 2]

# Calculate list stderr
list1_stderr = stderr(list1)
list2_stderr = stderr(list2)

# Print stderr values
print("Standard error for list1:", list1_stderr)
print("Standard error for list2:", list2_stderr)

# Check that the stderr values are correct
assert_equal(round(list1_stderr, 3), 0.632)
assert_equal(round(list2_stderr, 3), 1.275)
```

# Problem 2: Analyzing volcano data (*8 points*)

In this problem we'll utilize the functions generated in Problem 1 to calculate some statistics related to Holocene volcanoes from the [Smithsonian Institution's Global Volcanism Program](https://volcano.si.edu/).

**Notice**: Closely follow the instructions! For example, you should be sure to use **exactly** the same variable names mentioned in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions.

**Your score on this problem will be based on following criteria**:

- Reading and processing the [data file](data/GVP_Volcano_List_Holocene.csv)
- Cleaning the raw data as instructed
- Creating a new DataFrame of regional volcano statistics
- Including comments that explain what most lines in the code do
- Answering a couple questions at the end of the problem
- Uploading your notebook to your GitHub repository for this week's exercise

## Part 1: Reading the data file (*1.5 points*)

The first step in this exercise is to read in the [data file](data/GVP_Volcano_List_Holocene.csv) we're using. 

- For this, you should use the Pandas `read_csv()` funtion to read in the data file [data/GVP_Volcano_List_Holocene.csv](data/GVP_Volcano_List_Holocene.csv) into the variable `data`
    - Use the semicolon character `;` as the separator between columns
    - Skip the first row

<!-- grade_id: read_data, points: 0, grade: false, solution: true -->
```python
data = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_read_data_tests_1, points: 0.5, grade: true, solution: false -->
```python
# This test print should print the head and the column names
print(data.head(1))
print(data.columns)

```

<!-- grade_id: problem_2_read_data_tests_2, points: 0.5, grade: true, solution: false -->
```python
# This test print should print the first row of data
print(data.loc[0].values)

```

<!-- grade_id: problem_2_read_data_tests_3, points: 0.5, grade: true, solution: false -->
```python
# This test print should print the number of rows of data
print('The DataFrame contains', len(data),'rows.')

```

## Part 2: Preparing the data (*1 point*)

Before moving to creating our statistical functions, we need to clean up our data a bit.

- Create a subset of the `data` dataframe called `clean_data` that contains only volcanoes with elevations above sea level
- Drop rows from the `clean_data` DataFrame with NaNs in the 'Tectonic Setting' column

<!-- grade_id: clean_data, points: 0, grade: false, solution: true -->
```python
clean_data = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_clean_data_tests_1, points: 0.5, grade: true, solution: false -->
```python
# This test print should print the head
print(clean_data.head(1))

```

<!-- grade_id: problem_2_clean_data_tests_2, points: 0.5, grade: true, solution: false -->
```python
# This test print should print the last 5 tectonic settings
print(clean_data['Tectonic Setting'].tail(5))

```

## Part 3: Global volcano statistics (*1 point*)

With our clean data, we can now calculate some global statistical values for Holocene volcanoes.

- Calculate the mean elevation of all volcanoes in the `clean_data` DataFrame as variable `global_mean`
- Calculate the standard deviation in the elevation of all the volcanoes in the `clean_data` DataFrame as variable `global_stddev`
- Calculate the standard error in the elevation of all the volcanoes in the `clean_data` DataFrame as variable `global_stderr`

**Note**: You will need to use the `.values` attribute with the elevation data to calculate the requested values using the functions you created in Problem 1.

<!-- grade_id: global_stats, points: 0, grade: false, solution: true -->
```python
global_mean = None
global_stddev = None
global_stderr = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_global_stats_tests_1, points: 0.5, grade: true, solution: false -->
```python
# This should print the global mean volcano elevation
print('My global average volcano elevation:', global_mean, 'meters.\nThe expected value is approximately 1922 meters.')

```

<!-- grade_id: problem_2_global_stats_tests_2, points: 0.25, grade: true, solution: false -->
```python
# This should print the global standard deviation in volcano elevation
print('My global standard deviation in volcano elevation:', global_stddev, 'meters.\nThe expected value is approximately 1406 meters.')

```

<!-- grade_id: problem_2_global_stats_tests_3, points: 0.25, grade: true, solution: false -->
```python
# This should print the global standard deviation in volcano elevation
print('My global standard error in volcano elevation:', global_stderr, 'meters.\nThe expected value is approximately 39 meters.')

```

## Part 4: Compiling regional volcano statistics (*2 points*)

Using an approach similar to that in Part 3, we can now calculate some statistics on volcanoes in various geographic regions.

- Create a variable `regions` that contains the unique 'Region' values from the `clean_data` DataFrame
- Using a `for` loop, loop over each region and append its mean, standard deviation, and standard error in volcano elevation to the empty list variables `means`, `stddevs`, and `stderrs`
    - **Note**: In your `for` loop you should first extract a subset of data using `.loc` to select the elevation values for the region of interest
    - You will also want to convert your subset of elevations to an array of values using the `.values` attribute before using your functions

<!-- grade_id: regional_data, points: 0, grade: false, solution: true -->
```python
regions = None

means = []
stddevs = []
stderrs = []

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_regional_data_tests_1, points: 0.5, grade: true, solution: false -->
```python
# This should print the last value in the regions array
print('The last region is', regions[-1])

```

<!-- grade_id: problem_2_regional_data_tests_2, points: 0.5, grade: true, solution: false -->
```python
# This should print the last value in the means list
print('The last mean value is', means[-1])

```

<!-- grade_id: problem_2_regional_data_tests_3, points: 0.5, grade: true, solution: false -->
```python
# This should print the last value in the stddevs list
print('The last mean value is', stddevs[-1])

```

<!-- grade_id: problem_2_regional_data_tests_4, points: 0.5, grade: true, solution: false -->
```python
# This should print the last value in the stderrs list
print('The last stderr value is', stderrs[-1])

```

## Part 5: Creating a regional DataFrame (*1.5 points*)

The last step in this problem is to combine our new lists into a new Pandas DataFrame.

- Create a new DataFrame called `region_data` using `regions` as the index and the `means`, `stddevs`, and `stderrs` lists as the column data
    - **Hint**: When creating a DataFrame, column data can be specified in a dictionary using the `data` keyword (e.g., `data = {'Column heading': column_values}`), where `'Column heading'` would be the name of the column, and `column_values` would be the data in that column
    - **Hint**: The `index` parameter can be used to indicate the values to be used for the index

<!-- grade_id: region_dataframe, points: 0, grade: false, solution: true -->
```python
region_data = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_region_dataframe_tests_1, points: 0.5, grade: true, solution: false -->
```python
# This should print the last row in the region_data DataFrame
print(region_data.tail(1))

```

<!-- grade_id: problem_2_region_dataframe_tests_2, points: 0.5, grade: true, solution: false -->
```python
# This should print the shape of the DataFrame
print('DataFrame dimensions:', region_data.shape)

```

<!-- grade_id: problem_2_region_dataframe_tests_3, points: 0.5, grade: true, solution: false -->
```python
# This should print mean elevation for Alaska
print('Alaska mean volcano elevation:', region_data['Mean elevation'].loc['Alaska'])

```

## Part 6: Questions for Problem 2 (*1 point*)

1. How much do the standard deviation values vary among the regions? What do these values tell you about the volcano elevations in different regions?
2. Do you observe a large difference between the standard deviation and standard error values? Is it clear why you should always indicate whether reported values are standard deviations or standard errors?

<!-- grade_id: problem_2_questions, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

## Problem 3: Visualizing uncertainty (*9 points*)

In this problem we will continue to develop our Python mathematical and plotting skills by visualizing the regional volcano data using bar plots and the *normal distribution*. The normal distribution is a mathematical function with a bell shape, also known as the *Gauss function* or *Gaussian*. This function is centered on the mean value of a given set of values, where its value is greatest, and its value decreases away from the mean in what is hopefully a familiar form.

**Notice**: Closely follow the instructions! For example, you should be sure to use **exactly** the same variable names mentioned in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions.

**Your score on this problem will be based on following criteria**:

- Creating bar plots of the regional volcano data
- Properly defining a function for calulating normal distributions
- Creating a new DataFrame of regional volcano elevation distributions
- Including comments that explain what most lines in the code do
- Answering a couple questions at the end of the problem
- Uploading your notebook to your GitHub repository for this week's exercise

## Part 1: A bar plot of our regional volcano data (*1.5 points*)

We'll start visualizing our volcano data by creating a bar plot of the mean elevations along with their standard deviations.

- Create a bar plot of the mean elevations in the `region_data` DataFrame, including the standard deviation in elevation as an error bar
    - The data for the error bars can be assigned using the `yerr` parameter in the `.plot()` function
    - We suggest you use a slightly larger figure size of 12 by 8 inches, which can be set using the `figsize` parameter in the `.plot()` function
- Include a title and label on the y-axis
- Also add a black, dashed line spanning the width of the plot and indicating the global mean volcano elevation
    - Check the hints for this week's exercise about how to add this line
- Be sure to show a plot legend
- Finally, add a figure caption in the Markdown cell below the Python cell for your plot that describes the plot as if it were in a scientific journal article

<!-- grade_id: problem_3_plot_1, points: 1, grade: true, solution: true -->
```python
%matplotlib inline

ax = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_3_caption_1, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

## Part 2: Another bar plot of our regional volcano data (*1.5 points*)

Let's now produce a plot similar to that from Part 1, but using the *standard error* data for the error bars.

- Create a bar plot of the mean elevations in the `region_data` DataFrame, including the standard error in elevation as an error bar
    - The data for the error bars can be assigned using the `yerr` parameter in the `.plot()` function
    - We suggest you use a slightly larger figure size of 12 by 8 inches, which can be set using the `figsize` parameter in the `.plot()` function
- Include a title and label on the y-axis
- Also add a black, dashed line spanning the width of the plot and indicating the global mean volcano elevation
    - Check the hints for this week's exercise about how to add this line
- Be sure to show a plot legend
- Finally, add a figure caption in the Markdown cell below the Python cell for your plot that describes the plot as if it were in a scientific journal article

<!-- grade_id: problem_3_plot_2, points: 1, grade: true, solution: true -->
```python
%matplotlib inline

ax = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_3_caption_2, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

## Part 3: One more statistical function (*2 points*)

Your next task in this problem is to define the functions that will be needed to make plots of the normal distributions of the data from Problem 2.
We can start with the Gaussian function itself, which has a mathematical definition that is
\begin{equation}
  \Large
  G_{\bar{x}, \sigma_{x}}(x) = \frac{1}{\sigma_{x} \sqrt{2 \pi}} e^{-(x - \bar{x})^{2} / 2 \sigma_{x}^{2})}
\end{equation}

*Equation 4. The normal distribution*.

In this equation, $e$ is the exponential function, $x$ is the value for which the normal distribution is calculated (elevation in our case), $\bar{x}$ is the mean, and $\sigma_{x}$ is the standard deviation.
As we have already defined functions for the mean and standard deviation, you can use your earlier functions here.
**Note**: With the Gaussian function we are explicitly assuming here that the uncertainty is symmetric and follows a bell-shaped distribution about the mean.

For this part you should:

- Import NumPy to be able to use the exponential function and one for calculating the square root
- Define a function called `gaussian()` that you can use to calculate the normal distribution.
    - **Note**: You will have more than one parameter when defining this function.

<!-- grade_id: gaussian, points: 0, grade: false, solution: true -->
```python
import numpy as np

def gaussian(mean_value, stddev_value, x_array):
    """Returns the Gaussian distribution of a range of values."""
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_3_gaussian_tests_1, points: 1, grade: true, solution: false -->
```python
# This test should also work
print("The Gaussian value for this test should be 0.3989. My Gaussian value is:", gaussian(1, 1, [1]))

```

<!-- grade_id: problem_3_gaussian_tests_2, points: 1, grade: true, solution: false -->
```python
# This test should work
print("The Gaussian value for this test should be 0.1994. My Gaussian value is:", gaussian(2, 2, [2]))

```

## Part 4: Calculating normal distributions for our elevation data (*1 point*)

Your next task now that we have the functions defined is to calculate the normal distributions for each of the regional volcano elevations from Problem 2. You can use the mean and standard deviation values calculated in the `region_data` DataFrame for this.

- Define a NumPy array called `elevations` that goes from 0 to 8000 meters in 1001 steps
- Create a new empty Pandas DataFrame called `gaussians` with `elevations` as the index
- Use a `for` loop to loop over each region in the `regions` array defined back in Problem 2
    - For each region, calculate the normal distribution of volcano elevations using the regional mean elevation and standard deviation in elevation over the range of elevations in the `elevations` array
    - Save each normal distribution in a new column named for the region
        - As you may recall, you can easily add a new column by assigning values to a new column name in the DataFrame. For example, `df['New column'] = elevations` would assign the values from the `elevations` array to a new column in the DataFrame `df` called `'New column'`

<!-- grade_id: gaussians, points: 0, grade: false, solution: true -->
```python
# Array of elevations
elevations = None

# New DataFrame for normal distributions
gaussians = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_3_gaussians_tests_1, points: 0.5, grade: true, solution: false -->
```python
# This test print should work
print('The elevation at index 400 is', elevations[400], '. The expected value is 3200.0.')

```

<!-- grade_id: problem_3_gaussproblem_3_gaussians_tests_1ians_tests_2, points: 0.5, grade: true, solution: false -->
```python
# This should print the shape of the DataFrame
print('DataFrame dimensions:', gaussians.shape)

```

## Part 5: Plotting our volcano data, version 2.0 (*2 points*)

Finally, we can plot our normal distributions.

- Create a plot of the data in the `gaussians` DataFrame using the Pandas `.plot()` function
    - Include a title and again we suggest a figure size of 12 by 8 inches
- Label the x- and y-axes
    - The y-axis label should be 'Probability'
- Set the range for the y-axis to go from 0.0 to 0.0016
- Finally, add a figure caption in the Markdown cell below the Python cell for your plot that describes the plot as if it were in a scientific journal article

<!-- grade_id: problem_3_plot_3, points: 1.5, grade: true, solution: true -->
```python
ax = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_3_caption_3, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

## Part 6: Questions for Problem 3 (*1 point*)

1. How does the shape of the Gaussian distribution change for different values of the standard deviation?
2. Using this method of visualization, is it clear which regions have larger or smaller variations in average volcano elevation?

<!-- grade_id: problem_3_answers, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE